[1452146501] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452146501] SERVICE FLAPPING ALERT: Ownerslist22;Explorer;STARTED; Service appears to have started flapping (23.7% change >= 20.0% threshold)
[1452146504] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452146531] HOST ALERT: Ownerslist21;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452146568] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452146595] HOST ALERT: Ownerslist21;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452146791] SERVICE FLAPPING ALERT: Ownerslist22;NSClient++ Version;STOPPED; Service appears to have stopped flapping (4.1% change < 5.0% threshold)
[1452146895] SERVICE ALERT: Ownerslist14;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452146898] HOST ALERT: Ownerslist14;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.31)
[1452146962] HOST ALERT: Ownerslist14;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1452147005] SERVICE ALERT: Ownerslist14;NSClient++ Version;OK;SOFT;2;NSClient++ 0,4,1,73 2012-12-17
[1452147161] SERVICE ALERT: Ownerslist14;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452147183] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452147191] HOST ALERT: Ownerslist14;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1452147255] HOST ALERT: Ownerslist14;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1452147255] HOST FLAPPING ALERT: Ownerslist14;STARTED; Host appears to have started flapping (23.4% change > 20.0% threshold)
[1452147271] SERVICE ALERT: Ownerslist14;C:\ Drive Space;OK;SOFT;2;c:\ - total: 74.08 Gb - used: 30.29 Gb (41%) - free 43.79 Gb (59%)
[1452147331] SERVICE ALERT: Ownerslist14;Explorer;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452147334] HOST ALERT: Ownerslist14;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.31)
[1452147354] SERVICE ALERT: Ownerslist14;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.31 and port 12489: No route to host
[1452147357] HOST ALERT: Ownerslist14;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.31)
[1452147380] HOST ALERT: Ownerslist14;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.31)
[1452147405] SERVICE ALERT: Ownerslist14;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.31 and port 12489: No route to host
[1452147408] HOST ALERT: Ownerslist14;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.0.31)
[1452147444] SERVICE ALERT: Ownerslist14;Explorer;CRITICAL;SOFT;2;connect to address 192.168.0.31 and port 12489: No route to host
[1452147447] HOST ALERT: Ownerslist14;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.0.31)
[1452147474] SERVICE ALERT: Ownerslist14;CPU Load;CRITICAL;SOFT;2;connect to address 192.168.0.31 and port 12489: No route to host
[1452147477] HOST ALERT: Ownerslist14;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.0.31)
[1452147525] SERVICE ALERT: Ownerslist14;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.31 and port 12489: No route to host
[1452147528] HOST ALERT: Ownerslist14;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.0.31)
[1452147564] SERVICE ALERT: Ownerslist14;Explorer;CRITICAL;HARD;3;connect to address 192.168.0.31 and port 12489: No route to host
[1452147567] HOST ALERT: Ownerslist14;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.0.31)
[1452147574] SERVICE ALERT: Ownerslist14;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.31 and port 12489: No route to host
[1452147594] SERVICE ALERT: Ownerslist14;CPU Load;CRITICAL;HARD;3;connect to address 192.168.0.31 and port 12489: No route to host
[1452147597] HOST ALERT: Ownerslist14;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.0.31)
[1452147597] HOST FLAPPING ALERT: Ownerslist14;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1452147608] SERVICE ALERT: Ownerslist14;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.31 and port 12489: No route to host
[1452147645] SERVICE ALERT: Ownerslist14;Memory Usage;CRITICAL;HARD;3;connect to address 192.168.0.31 and port 12489: No route to host
[1452147648] HOST ALERT: Ownerslist14;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.0.31)
[1452147694] SERVICE ALERT: Ownerslist14;Uptime;CRITICAL;SOFT;2;connect to address 192.168.0.31 and port 12489: No route to host
[1452147728] SERVICE ALERT: Ownerslist14;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.31 and port 12489: No route to host
[1452147783] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452147814] SERVICE ALERT: Ownerslist14;Uptime;CRITICAL;HARD;3;connect to address 192.168.0.31 and port 12489: No route to host
[1452147824] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452147848] SERVICE ALERT: Ownerslist14;NSClient++ Version;CRITICAL;HARD;3;connect to address 192.168.0.31 and port 12489: No route to host
[1452147874] SERVICE ALERT: Ownerslist14;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.31 and port 12489: No route to host
[1452147994] SERVICE ALERT: Ownerslist14;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.31 and port 12489: No route to host
[1452148071] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452148114] SERVICE ALERT: Ownerslist14;C:\ Drive Space;CRITICAL;HARD;3;connect to address 192.168.0.31 and port 12489: No route to host
[1452148191] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452148304] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452148311] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452148368] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452148791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452148904] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452148932] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452148934] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452148934] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452148968] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.89 ms
[1452148968] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1452149042] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452149046] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1452149051] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452149054] HOST ALERT: Ownerslist14;UP;HARD;10;PING OK - Packet loss = 28%, RTA = 78.33 ms
[1452149214] HOST ALERT: Ownerslist14;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1452149278] HOST ALERT: Ownerslist14;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.87 ms
[1452149311] SERVICE ALERT: Ownerslist14;C:\ Drive Space;OK;HARD;3;c:\ - total: 74.08 Gb - used: 30.29 Gb (41%) - free 43.79 Gb (59%)
[1452149324] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452149327] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452149374] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452149374] HOST ALERT: Ownerslist14;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.31)
[1452149374] HOST FLAPPING ALERT: Ownerslist14;STARTED; Host appears to have started flapping (23.3% change > 20.0% threshold)
[1452149377] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452149394] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452149397] HOST ALERT: Ownerslist14;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.31)
[1452149397] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452149441] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2580.18 MB (42%) - free: 3534.17 MB (58%)
[1452149445] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452149445] HOST FLAPPING ALERT: Ownerslist15;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1452149448] HOST ALERT: Ownerslist14;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.31)
[1452149491] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;2;System Uptime - 17 day(s) 21 hour(s) 19 minute(s)
[1452149511] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 51% (5 min average)
[1452149511] HOST ALERT: Ownerslist14;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.0.31)
[1452149574] HOST ALERT: Ownerslist14;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.0.31)
[1452149617] HOST ALERT: Ownerslist14;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.0.31)
[1452149651] HOST ALERT: Ownerslist14;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.0.31)
[1452149714] HOST ALERT: Ownerslist14;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.0.31)
[1452149777] HOST ALERT: Ownerslist14;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.0.31)
[1452149781] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452149784] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452149797] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452149805] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452149861] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452149891] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452149914] SERVICE ALERT: Ownerslist14;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.31 and port 12489: No route to host
[1452149917] HOST ALERT: Ownerslist14;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.0.31)
[1452149967] HOST FLAPPING ALERT: Ownerslist14;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1452149991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452150034] SERVICE ALERT: Ownerslist14;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.31 and port 12489: No route to host
[1452150101] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452150154] SERVICE ALERT: Ownerslist14;C:\ Drive Space;CRITICAL;HARD;3;connect to address 192.168.0.31 and port 12489: No route to host
[1452150252] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452150372] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452150375] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452150407] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452150407] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452150416] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452150427] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 1.03 ms
[1452150438] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452150482] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452150486] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452150521] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.61 Gb (49%) - free 37.38 Gb (51%)
[1452150533] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1541.60 MB (25%) - free: 4572.75 MB (75%)
[1452150612] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452150628] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452150651] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452150654] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452150685] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452150701] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452150711] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452150722] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452150738] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452150761] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2589.35 MB (42%) - free: 3525.00 MB (58%)
[1452150833] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452150836] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452150861] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452150864] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452150900] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452150928] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452150941] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 18 day(s) 22 hour(s) 44 minute(s)
[1452150971] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452151092] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452151212] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452151230] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452151242] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452151285] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.36 ms
[1452151294] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452151322] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452151383] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452151731] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452151743] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452151761] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452151791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452151825] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452151841] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.62 Gb (49%) - free 37.37 Gb (51%)
[1452151853] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1542.84 MB (25%) - free: 4571.51 MB (75%)
[1452151892] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452151892] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452151901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452151916] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452151932] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452151935] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452151941] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452151983] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452152011] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;2;System Uptime - 17 day(s) 22 hour(s) 1 minute(s)
[1452152042] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452152046] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452152058] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452152151] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452152155] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452152219] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.78 ms
[1452152261] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 18 day(s) 23 hour(s) 6 minute(s)
[1452152321] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452152431] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.21 Gb (18%) - free 121.26 Gb (82%)
[1452152494] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452152501] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452152511] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452152514] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452152517] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452152525] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452152565] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452152631] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452152631] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452152642] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452152791] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452152825] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1452152865] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452152991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452153391] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452153407] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452153465] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452153511] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452153531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452153631] HOST FLAPPING ALERT: Ownerslist15;STOPPED; Host appears to have stopped flapping (4.3% change < 5.0% threshold)
[1452153631] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452153651] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452153654] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452153654] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452153695] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452153701] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452153761] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.64 Gb (50%) - free 37.35 Gb (50%)
[1452153773] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1564.38 MB (26%) - free: 4549.96 MB (74%)
[1452153852] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452153855] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452153861] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452153962] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452153966] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452153978] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452154071] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452154074] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452154095] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452154131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452154181] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 18 day(s) 23 hour(s) 38 minute(s)
[1452154191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452154301] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452154701] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452154704] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452154768] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.65 ms
[1452154811] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452154831] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452154904] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452154965] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452155031] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452155151] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452155271] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452155271] SERVICE FLAPPING ALERT: Ownerslist22;Uptime;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1452155391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452155504] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452155565] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452156031] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452156151] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452156171] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452156174] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452156174] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452156271] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452156271] SERVICE FLAPPING ALERT: Ownerslist22;NSClient++ Version;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1452156281] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.66 Gb (50%) - free 37.33 Gb (50%)
[1452156285] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452156293] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1571.72 MB (26%) - free: 4542.63 MB (74%)
[1452156372] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452156375] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452156381] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452156482] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 21% (5 min average)
[1452156486] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1452156498] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452156591] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452156594] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452156615] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.87 ms
[1452156701] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 0 hour(s) 20 minute(s)
[1452157025] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1452157086] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1452157221] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452157224] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452157241] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452157244] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452157285] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452157288] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452157301] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452157331] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452157351] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.21 Gb (18%) - free 121.26 Gb (82%)
[1452157421] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452157424] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452157434] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452157445] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452157448] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452157512] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1452157512] HOST FLAPPING ALERT: Ownerslist15;STARTED; Host appears to have started flapping (22.1% change > 20.0% threshold)
[1452157531] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;2;System Uptime - 17 day(s) 23 hour(s) 33 minute(s)
[1452157551] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452157562] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452157600] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452157664] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452157824] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452157885] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452157901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452157971] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452158001] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452158065] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.96 ms
[1452158081] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2596.59 MB (42%) - free: 3517.76 MB (58%)
[1452158161] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452158172] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452158181] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452158191] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452158255] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1452158271] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452158271] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 52% (5 min average)
[1452158282] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452158291] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452158501] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452158691] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452158691] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452158694] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452158694] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452158694] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452158735] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452158801] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2606.15 MB (43%) - free: 3508.20 MB (57%)
[1452158801] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.56 Gb (49%) - free 37.43 Gb (51%)
[1452158805] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452158813] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1590.98 MB (26%) - free: 4523.37 MB (74%)
[1452158881] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452158884] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452158884] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452158892] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452158894] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452158895] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452158901] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452158991] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452158991] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452158995] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1452159002] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452159002] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452159006] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452159011] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452159018] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452159094] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452159104] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452159104] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452159135] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.94 ms
[1452159155] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.89 ms
[1452159221] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 1 hour(s) 2 minute(s)
[1452159467] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452159471] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452159531] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 601.38 ms
[1452159591] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452159694] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452159701] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452159704] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452159734] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452159737] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452159755] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.48 ms
[1452159801] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.04 ms
[1452159851] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452159991] SERVICE ALERT: new-pc;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452160053] SERVICE ALERT: new-pc;Explorer;CRITICAL;SOFT;1;connect to address 192.168.0.205 and port 12489: No route to host
[1452160067] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452160104] SERVICE ALERT: new-pc;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452160131] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452160131] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452160173] SERVICE ALERT: new-pc;Explorer;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452160191] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452160195] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1452160224] SERVICE ALERT: new-pc;NSClient++ Version;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452160293] SERVICE ALERT: new-pc;Explorer;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452160294] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452160301] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452160306] SERVICE ALERT: new-pc;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.205 and port 12489: No route to host
[1452160355] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452160426] SERVICE ALERT: new-pc;Uptime;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452160475] SERVICE ALERT: new-pc;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.205 and port 12489: No route to host
[1452160485] SERVICE ALERT: new-pc;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.205 and port 12489: No route to host
[1452160494] SERVICE ALERT: new-pc;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.205 and port 12489: No route to host
[1452160546] SERVICE ALERT: new-pc;Uptime;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452160595] SERVICE ALERT: new-pc;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452160605] SERVICE ALERT: new-pc;CPU Load;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452160614] SERVICE ALERT: new-pc;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452160667] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452160715] SERVICE ALERT: new-pc;C:\ Drive Space;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452160725] SERVICE ALERT: new-pc;CPU Load;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452160731] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452160734] SERVICE ALERT: new-pc;Memory Usage;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452160814] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452160817] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452160880] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452160901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452160931] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452160935] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1452161207] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452161207] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452161211] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452161214] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452161216] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452161227] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 58%, RTA = 0.90 ms
[1452161277] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452161321] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2600.06 MB (43%) - free: 3514.29 MB (57%)
[1452161321] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.58 Gb (49%) - free 37.41 Gb (51%)
[1452161325] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.76 ms
[1452161333] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1590.22 MB (26%) - free: 4524.13 MB (74%)
[1452161391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452161401] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452161404] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452161404] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452161412] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452161415] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452161419] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452161487] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452161495] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452161511] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 56% (5 min average)
[1452161515] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.88 ms
[1452161522] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452161522] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452161536] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452161631] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452161634] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452161655] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452161741] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 1 hour(s) 44 minute(s)
[1452161827] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452161890] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452161925] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.15 ms
[1452161991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452162261] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452162264] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452162328] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452162371] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452162591] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452162640] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452162698] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.95 ms
[1452162997] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452163060] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452163124] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452163304] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452163329] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.81 ms
[1452163383] SERVICE FLAPPING ALERT: Ownerslist22;Memory Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452163600] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452163663] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452163725] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.06 ms
[1452163731] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452163734] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452163734] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452163791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452163791] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452163841] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.61 Gb (49%) - free 37.38 Gb (51%)
[1452163845] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.90 ms
[1452163853] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1588.60 MB (26%) - free: 4525.75 MB (74%)
[1452163894] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452163904] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452163911] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452163915] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452163925] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452163928] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452163939] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452164042] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452164046] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452164056] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452164151] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452164227] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452164261] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 2 hour(s) 26 minute(s)
[1452164290] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452164325] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 3.11 ms
[1452164391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452164457] SERVICE ALERT: Ownerslist21;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452164494] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452164511] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452164514] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452164518] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 591.67 ms
[1452164567] SERVICE ALERT: Ownerslist21;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.71 Gb - used: 39.49 Gb (27%) - free 109.22 Gb (73%)
[1452164631] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452164641] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452164671] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452164735] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1452164751] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452164751] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;3;CPU Load 50% (5 min average)
[1452164781] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452164784] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452164848] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.81 ms
[1452164891] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452165043] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452165087] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452165125] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.68 ms
[1452165351] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452165407] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452165437] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452165465] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1452165500] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452165525] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.32 ms
[1452165531] SERVICE FLAPPING ALERT: Ownerslist22;C:\ Drive Space;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452165694] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452165725] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.14 ms
[1452165783] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452166037] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452166101] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452166101] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452166104] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452166168] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 6.56 ms
[1452166211] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452166304] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452166368] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1452166383] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452166551] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452166671] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452166894] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452166926] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452167071] SERVICE FLAPPING ALERT: Ownerslist22;NSClient++ Version;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452167238] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452167271] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452167301] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452167325] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.00 ms
[1452167391] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452167494] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452167511] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452167525] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452167642] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452167825] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.05 seconds)
[1452167837] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452167865] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1452167901] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 602.80 ms
[1452167991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452168021] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452168094] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452168125] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.85 ms
[1452168131] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452168264] SERVICE ALERT: Ownerslist21;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452168374] SERVICE ALERT: Ownerslist21;CPU Load;OK;SOFT;2;CPU Load 5% (5 min average)
[1452168407] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452168437] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452168465] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452168465] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (21.3% change > 20.0% threshold)
[1452168500] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452168525] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.13 ms
[1452168651] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452168654] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452168654] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452168695] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452168761] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.72 Gb (50%) - free 37.27 Gb (50%)
[1452168773] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1585.59 MB (26%) - free: 4528.76 MB (74%)
[1452168852] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452168855] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452168859] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452168962] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452168966] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452168976] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452169071] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452169074] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452169138] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1452169181] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 3 hour(s) 48 minute(s)
[1452169294] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452169301] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452169326] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452169638] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452169701] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452169725] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.24 ms
[1452169894] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452169904] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452169925] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.63 ms
[1452169935] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.35 ms
[1452170237] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452170300] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452170325] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.90 ms
[1452170494] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452170504] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452170511] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452170525] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452170534] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452170537] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452170565] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452170651] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452170837] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452170900] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452170925] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.99 ms
[1452171094] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452171101] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452171111] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452171125] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452171171] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452171174] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452171174] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452171231] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452171255] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452171281] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.68 Gb (50%) - free 37.31 Gb (50%)
[1452171293] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1592.13 MB (26%) - free: 4522.22 MB (74%)
[1452171372] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452171375] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452171379] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452171437] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452171482] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452171486] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452171496] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452171500] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452171525] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452171591] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452171591] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452171594] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452171658] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.81 ms
[1452171694] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452171701] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452171701] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 4 hour(s) 30 minute(s)
[1452171725] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452171831] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452171951] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452172007] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452172037] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452172071] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452172071] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.99 ms
[1452172100] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452172125] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.05 ms
[1452172191] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452172294] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452172301] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452172325] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1452172606] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452172637] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452172670] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1452172701] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 201.11 ms
[1452172791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452172894] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452172925] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 4.57 ms
[1452173061] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452173064] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452173085] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452173141] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452173171] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452173261] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452173264] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452173324] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452173327] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452173371] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 3 hour(s) 57 minute(s)
[1452173375] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452173391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452173441] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2624.14 MB (43%) - free: 3490.21 MB (57%)
[1452173494] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452173524] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452173527] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452173534] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452173554] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452173557] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452173621] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452173641] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452173651] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452173681] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452173690] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452173691] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452173694] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452173694] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452173753] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452173753] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452173775] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452173791] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;3;c:\ - total: 148.47 Gb - used: 27.25 Gb (18%) - free 121.22 Gb (82%)
[1452173795] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1452173801] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.58 Gb (49%) - free 37.41 Gb (51%)
[1452173813] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1586.54 MB (26%) - free: 4527.81 MB (74%)
[1452173871] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452173892] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452173895] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452173899] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452173981] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452173991] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452174002] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452174006] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452174016] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452174101] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452174104] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452174104] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452174118] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452174181] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452174211] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 4 hour(s) 11 minute(s)
[1452174215] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.11 ms
[1452174221] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 5 hour(s) 12 minute(s)
[1452174225] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452174401] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452174435] HOST ALERT: Ownerslist21;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452174456] HOST ALERT: Ownerslist21;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452174471] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452174521] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452174537] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452174544] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452174601] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.91 ms
[1452174607] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452174631] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;3;c:\ - total: 148.47 Gb - used: 27.26 Gb (18%) - free 121.21 Gb (82%)
[1452174635] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452174704] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452174768] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452174821] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452174824] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452174845] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452174941] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452174944] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452174984] SERVICE ALERT: Ownerslist21;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452175007] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452175051] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 4 hour(s) 25 minute(s)
[1452175055] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.69 ms
[1452175071] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452175094] SERVICE ALERT: Ownerslist21;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1452175191] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452175191] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452175241] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452175244] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452175244] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452175285] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452175311] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452175361] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452175364] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.252 and port 12489: No route to host
[1452175364] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452175444] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452175447] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452175454] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452175471] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;3;c:\ - total: 148.47 Gb - used: 27.27 Gb (18%) - free 121.20 Gb (82%)
[1452175475] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452175481] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;3;Memory usage: total:6114.35 MB - used: 2635.34 MB (43%) - free: 3479.01 MB (57%)
[1452175571] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452175574] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.252 and port 12489: No route to host
[1452175574] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452175581] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452175584] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452175605] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452175654] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452175657] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452175675] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1452175681] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452175691] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1452175691] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452175781] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452175784] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452175791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452175847] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452175885] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452175891] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 4 hour(s) 39 minute(s)
[1452176081] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452176084] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452176084] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452176201] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452176207] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452176207] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452176211] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452176216] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452176226] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452176228] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 61%, RTA = 0.94 ms
[1452176274] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452176277] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452176284] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452176294] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452176299] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 16%, RTA = 127.59 ms
[1452176311] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;3;c:\ - total: 148.47 Gb - used: 27.28 Gb (18%) - free 121.20 Gb (82%)
[1452176321] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;3;Memory usage: total:6114.35 MB - used: 2603.45 MB (43%) - free: 3510.90 MB (57%)
[1452176321] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.60 Gb (49%) - free 37.39 Gb (51%)
[1452176333] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1590.67 MB (26%) - free: 4523.68 MB (74%)
[1452176391] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452176411] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452176412] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452176415] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452176419] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452176421] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452176441] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452176488] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1452176495] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452176521] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452176522] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452176531] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1452176536] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452176631] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452176634] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452176698] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1452176741] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 5 hour(s) 54 minute(s)
[1452176862] SERVICE ALERT: Ownerslist21;W35VC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1452176991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452177028] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452177092] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.79 ms
[1452177094] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452177094] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452177125] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452177211] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;2;System Uptime - 18 day(s) 5 hour(s) 1 minute(s)
[1452177437] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452177452] SERVICE ALERT: Ownerslist21;W35VC;UNKNOWN;HARD;3;Failed to open service W3SVC: 1060: The specified service does not exist as an installed service.
[1452177500] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452177514] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452177517] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452177524] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452177595] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452177641] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452177644] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.252 and port 12489: No route to host
[1452177644] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452177687] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452177701] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452177724] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452177727] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452177734] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452177751] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;3;c:\ - total: 148.47 Gb - used: 27.17 Gb (18%) - free 121.31 Gb (82%)
[1452177755] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452177761] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;3;Memory usage: total:6114.35 MB - used: 2612.36 MB (43%) - free: 3501.99 MB (57%)
[1452177851] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452177854] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.252 and port 12489: No route to host
[1452177854] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452177855] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452177961] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452177965] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452177971] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1452177978] HOST ALERT: Ownerslist14;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452178161] SERVICE ALERT: Ownerslist14;Explorer;OK;HARD;3;explorer.exe: Running
[1452178191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452178191] SERVICE ALERT: Ownerslist14;CPU Load;OK;HARD;3;CPU Load 4% (5 min average)
[1452178241] SERVICE ALERT: Ownerslist14;Memory Usage;OK;HARD;3;Memory usage: total:4719.34 MB - used: 1513.45 MB (32%) - free: 3205.88 MB (68%)
[1452178277] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452178287] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452178295] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452178350] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452178351] SERVICE ALERT: Ownerslist14;C:\ Drive Space;OK;HARD;3;c:\ - total: 74.08 Gb - used: 30.29 Gb (41%) - free 43.79 Gb (59%)
[1452178354] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452178364] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452178370] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 28%, RTA = 114.31 ms
[1452178411] SERVICE ALERT: Ownerslist14;Uptime;OK;HARD;3;System Uptime - 2 day(s) 18 hour(s) 26 minute(s)
[1452178445] SERVICE ALERT: Ownerslist14;NSClient++ Version;OK;HARD;3;NSClient++ 0,4,1,73 2012-12-17
[1452178455] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452178471] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.17 Gb (18%) - free 121.30 Gb (82%)
[1452178491] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452178503] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452178564] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452178564] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452178574] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452178575] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452178601] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;3;Memory usage: total:6114.35 MB - used: 2614.63 MB (43%) - free: 3499.72 MB (57%)
[1452178605] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 3.29 ms
[1452178681] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452178694] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.252 and port 12489: No route to host
[1452178695] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452178697] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452178760] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452178795] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452178811] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1452178894] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452178901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452178957] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452179015] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.38 ms
[1452179291] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452179294] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452179357] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452179391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452179395] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452179401] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452179494] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452179557] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452179615] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.28 ms
[1452179663] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1452179783] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;SOFT;2;CPU Load 92% (5 min average)
[1452179895] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452179903] SERVICE ALERT: Ownerslist24;CPU Load;WARNING;HARD;3;CPU Load 89% (5 min average)
[1452179927] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452179991] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.91 ms
[1452180094] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452180157] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452180215] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452180495] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452180503] SERVICE ALERT: Ownerslist24;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452180527] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452180590] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452180594] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452180615] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452180615] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.18 ms
[1452180721] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452180724] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452180735] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452180735] SERVICE FLAPPING ALERT: Ownerslist22;CPU Load;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1452180787] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452180815] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.57 ms
[1452180831] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;3;CPU Load 51% (5 min average)
[1452181011] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452181131] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452181131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452181134] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452181197] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452181204] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452181215] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.91 ms
[1452181241] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;3;Memory usage: total:6114.35 MB - used: 2606.33 MB (43%) - free: 3508.01 MB (57%)
[1452181331] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452181354] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452181415] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452181441] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452181703] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;SOFT;1;CPU Load 95% (5 min average)
[1452181731] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452181731] SERVICE ALERT: new-pc;CPU Load;WARNING;HARD;3;No data was received from host!
[1452181735] SERVICE ALERT: new-pc;Memory Usage;OK;HARD;3;Memory usage: total:8188.76 MB - used: 2152.58 MB (26%) - free: 6036.18 MB (74%)
[1452181745] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452181808] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452181821] SERVICE ALERT: new-pc;NSClient++ Version;OK;HARD;3;NSClient++ 0,4,1,73 2012-12-17
[1452181822] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;SOFT;2;CPU Load 93% (5 min average)
[1452181845] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1452181889] SERVICE ALERT: new-pc;Explorer;OK;HARD;3;explorer.exe: Running
[1452181942] SERVICE ALERT: Ownerslist24;CPU Load;OK;SOFT;3;CPU Load 6% (5 min average)
[1452182021] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452182141] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452182141] SERVICE ALERT: new-pc;Uptime;OK;HARD;3;System Uptime - 1 day(s) 5 hour(s) 20 minute(s)
[1452182144] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452182207] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452182251] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 6 hour(s) 25 minute(s)
[1452182255] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452182311] SERVICE ALERT: new-pc;C:\ Drive Space;OK;HARD;3;c:\ - total: 465.66 Gb - used: 56.35 Gb (12%) - free 409.31 Gb (88%)
[1452182322] SERVICE ALERT: new-pc;CPU Load;OK;HARD;3;CPU Load 19% (5 min average)
[1452182421] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452182424] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452182447] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452182531] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452182583] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452182641] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452182651] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452182671] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452182751] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452182771] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452182788] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452182851] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452182851] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452182881] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452182885] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452182981] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452182984] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452183044] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452183044] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452183085] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452183091] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 6 hour(s) 39 minute(s)
[1452183101] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452183142] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1452183161] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2614.20 MB (43%) - free: 3500.15 MB (57%)
[1452183183] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452183262] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;SOFT;2;CPU Load 96% (5 min average)
[1452183271] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452183281] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452183382] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;HARD;3;CPU Load 95% (5 min average)
[1452183391] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452183401] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452183426] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452183484] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452183487] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452183511] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;3;c:\ - total: 148.47 Gb - used: 27.22 Gb (18%) - free 121.25 Gb (82%)
[1452183515] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452183611] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452183614] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452183677] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452183694] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452183697] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452183721] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452183725] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452183811] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;2;System Uptime - 18 day(s) 6 hour(s) 51 minute(s)
[1452183982] SERVICE ALERT: Ownerslist24;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452183982] SERVICE FLAPPING ALERT: Ownerslist24;CPU Load;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1452183991] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452184111] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452184191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452184311] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452184589] HOST FLAPPING ALERT: Ownerslist15;STOPPED; Host appears to have stopped flapping (4.4% change < 5.0% threshold)
[1452184901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452184911] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452185031] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452185291] SERVICE ALERT: Ownerslist13;W35VC;OK;HARD;3;W3SVC: Started
[1452185295] HOST ALERT: Ownerslist13;UP;HARD;10;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1452185311] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452185351] SERVICE ALERT: Ownerslist13;Explorer;OK;HARD;3;explorer.exe: Running
[1452185391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452185431] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452185441] SERVICE ALERT: Ownerslist13;Memory Usage;OK;HARD;3;Memory usage: total:6114.35 MB - used: 1318.41 MB (22%) - free: 4795.94 MB (78%)
[1452185524] SERVICE ALERT: Ownerslist13;NSClient++ Version;OK;HARD;3;NSClient++ 0,4,1,73 2012-12-17
[1452185551] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;3;(No output returned from plugin)
[1452185611] SERVICE ALERT: Ownerslist13;Uptime;OK;HARD;3;System Uptime - 2 day(s) 20 hour(s) 25 minute(s)
[1452185783] SERVICE ALERT: Ownerslist13;C:\ Drive Space;OK;HARD;3;c:\ - total: 73.99 Gb - used: 20.46 Gb (28%) - free 53.53 Gb (72%)
[1452185871] SERVICE ALERT: Ownerslist13;CPU Load;OK;HARD;3;CPU Load 3% (5 min average)
[1452186531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452186751] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452186871] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452186991] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452187131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452187131] SERVICE FLAPPING ALERT: Ownerslist22;C:\ Drive Space;STARTED; Service appears to have started flapping (22.4% change >= 20.0% threshold)
[1452187383] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452187791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452187983] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452187983] SERVICE FLAPPING ALERT: Ownerslist22;Memory Usage;STARTED; Service appears to have started flapping (22.6% change >= 20.0% threshold)
[1452188331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452188391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452188931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452189133] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452189733] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452189831] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452189853] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452189951] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452189973] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452190191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452190551] SERVICE FLAPPING ALERT: Ownerslist22;Uptime;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1452190791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452190983] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452191583] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452191751] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452191871] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452191911] SERVICE ALERT: Ownerslist27;CPU Load;CRITICAL;SOFT;1;CPU Load 100% (5 min average)
[1452191991] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452191991] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452192031] SERVICE ALERT: Ownerslist27;CPU Load;CRITICAL;SOFT;2;CPU Load 100% (5 min average)
[1452192151] SERVICE ALERT: Ownerslist27;CPU Load;CRITICAL;HARD;3;CPU Load 92% (5 min average)
[1452192391] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452192531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452192701] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452192751] SERVICE ALERT: Ownerslist27;CPU Load;OK;HARD;3;CPU Load 36% (5 min average)
[1452192991] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452193111] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452193114] SERVICE ALERT: Ownerslist21;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452193131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452193144] HOST ALERT: Ownerslist21;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452193191] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452193208] HOST ALERT: Ownerslist21;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452193224] SERVICE ALERT: Ownerslist21;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452193231] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;3;(No output returned from plugin)
[1452193301] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452193311] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452193431] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452193561] SERVICE ALERT: Ownerslist16;CPU Load;CRITICAL;SOFT;1;CPU Load 90% (5 min average)
[1452193681] SERVICE ALERT: Ownerslist16;CPU Load;OK;SOFT;2;CPU Load 10% (5 min average)
[1452193831] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452193901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452193951] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452194031] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452194071] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;3;(No output returned from plugin)
[1452194391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452194631] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452194671] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452194751] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452194791] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452194921] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452194931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452194951] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.05 seconds)
[1452194991] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452195015] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452195031] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.43 Gb (18%) - free 121.05 Gb (82%)
[1452195051] SERVICE ALERT: localhost;Root Partition;CRITICAL;HARD;4;DISK CRITICAL - free space: / 738 MB (10% inode=95%):
[1452195131] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452195141] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452195161] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1452195171] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452195185] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 66%, RTA = 1.63 ms
[1452195185] HOST FLAPPING ALERT: Ownerslist15;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1452195241] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452195251] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452195281] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2653.98 MB (43%) - free: 3460.37 MB (57%)
[1452195382] SERVICE FLAPPING ALERT: Ownerslist24;CPU Load;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452195391] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452195435] HOST ALERT: Ownerslist21;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1452195455] HOST ALERT: Ownerslist21;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452195511] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452195531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452195631] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;3;(No output returned from plugin)
[1452195951] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452196071] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452196131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452196285] HOST FLAPPING ALERT: Ownerslist15;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452196573] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452196901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452197572] SERVICE ALERT: Ownerslist21;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452197583] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452197682] SERVICE ALERT: Ownerslist21;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.71 Gb - used: 39.67 Gb (27%) - free 109.03 Gb (73%)
[1452197991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452198031] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452198151] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452198373] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452198418] HOST ALERT: Ownerslist24;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.16)
[1452198482] HOST ALERT: Ownerslist24;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1452198493] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452198531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452198613] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452198751] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452198871] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452198991] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452199071] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452199191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452199191] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452199301] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452199383] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452199551] SERVICE ALERT: Ownerslist16;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452199554] HOST ALERT: Ownerslist16;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.123)
[1452199582] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;SOFT;1;CPU Load 97% (5 min average)
[1452199596] SERVICE ALERT: Ownerslist16;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.123 and port 12489: No route to host
[1452199599] HOST ALERT: Ownerslist16;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.123)
[1452199662] SERVICE ALERT: Ownerslist16;Uptime;CRITICAL;SOFT;2;connect to address 192.168.0.123 and port 12489: No route to host
[1452199662] HOST ALERT: Ownerslist16;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.123)
[1452199684] SERVICE ALERT: Ownerslist16;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.123 and port 12489: No route to host
[1452199687] HOST ALERT: Ownerslist16;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.0.123)
[1452199702] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;SOFT;2;CPU Load 93% (5 min average)
[1452199716] SERVICE ALERT: Ownerslist16;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.123 and port 12489: No route to host
[1452199719] HOST ALERT: Ownerslist16;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.0.123)
[1452199785] SERVICE ALERT: Ownerslist16;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452199786] HOST ALERT: Ownerslist16;UP;SOFT;6;PING OK - Packet loss = 37%, RTA = 2.65 ms
[1452199801] SERVICE ALERT: Ownerslist16;CPU Load;OK;SOFT;2;CPU Load 7% (5 min average)
[1452199822] SERVICE ALERT: Ownerslist24;CPU Load;CRITICAL;HARD;3;CPU Load 96% (5 min average)
[1452199833] SERVICE ALERT: Ownerslist16;C:\ Drive Space;OK;SOFT;3;c:\ - total: 148.71 Gb - used: 47.41 Gb (32%) - free 101.29 Gb (68%)
[1452199901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452199981] SERVICE ALERT: Ownerslist16;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452199984] HOST ALERT: Ownerslist16;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.123)
[1452200047] HOST ALERT: Ownerslist16;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.123)
[1452200094] SERVICE ALERT: Ownerslist16;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.123 and port 12489: No route to host
[1452200097] HOST ALERT: Ownerslist16;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.123)
[1452200160] HOST ALERT: Ownerslist16;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.0.123)
[1452200214] SERVICE ALERT: Ownerslist16;Memory Usage;CRITICAL;HARD;3;connect to address 192.168.0.123 and port 12489: No route to host
[1452200217] HOST ALERT: Ownerslist16;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.0.123)
[1452200280] HOST ALERT: Ownerslist16;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.0.123)
[1452200331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452200343] HOST ALERT: Ownerslist16;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.0.123)
[1452200384] SERVICE ALERT: Ownerslist16;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.123 and port 12489: No route to host
[1452200387] HOST ALERT: Ownerslist16;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.0.123)
[1452200404] SERVICE ALERT: Ownerslist16;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.123 and port 12489: No route to host
[1452200407] HOST ALERT: Ownerslist16;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.0.123)
[1452200414] SERVICE ALERT: Ownerslist16;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.123 and port 12489: No route to host
[1452200436] SERVICE ALERT: Ownerslist16;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.123 and port 12489: No route to host
[1452200439] HOST ALERT: Ownerslist16;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.0.123)
[1452200454] SERVICE ALERT: Ownerslist16;W35VC;CRITICAL;HARD;3;connect to address 192.168.0.123 and port 12489: No route to host
[1452200504] SERVICE ALERT: Ownerslist16;Uptime;CRITICAL;SOFT;2;connect to address 192.168.0.123 and port 12489: No route to host
[1452200524] SERVICE ALERT: Ownerslist16;CPU Load;CRITICAL;SOFT;2;connect to address 192.168.0.123 and port 12489: No route to host
[1452200534] SERVICE ALERT: Ownerslist16;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.123 and port 12489: No route to host
[1452200556] SERVICE ALERT: Ownerslist16;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.123 and port 12489: No route to host
[1452200624] SERVICE ALERT: Ownerslist16;Uptime;CRITICAL;HARD;3;connect to address 192.168.0.123 and port 12489: No route to host
[1452200644] SERVICE ALERT: Ownerslist16;CPU Load;CRITICAL;HARD;3;connect to address 192.168.0.123 and port 12489: No route to host
[1452200654] SERVICE ALERT: Ownerslist16;NSClient++ Version;CRITICAL;HARD;3;connect to address 192.168.0.123 and port 12489: No route to host
[1452200676] SERVICE ALERT: Ownerslist16;C:\ Drive Space;CRITICAL;HARD;3;connect to address 192.168.0.123 and port 12489: No route to host
[1452201022] SERVICE ALERT: Ownerslist24;CPU Load;OK;HARD;3;CPU Load 10% (5 min average)
[1452201151] SERVICE ALERT: Ownerslist27;CPU Load;WARNING;SOFT;1;CPU Load 83% (5 min average)
[1452201271] SERVICE ALERT: Ownerslist27;CPU Load;CRITICAL;SOFT;2;CPU Load 99% (5 min average)
[1452201391] SERVICE ALERT: Ownerslist27;CPU Load;CRITICAL;HARD;3;CPU Load 99% (5 min average)
[1452201531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452202131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452202731] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452203761] SERVICE ALERT: Ownerslist21;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452203791] HOST ALERT: Ownerslist21;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452203791] SERVICE ALERT: Ownerslist27;CPU Load;OK;HARD;3;CPU Load 53% (5 min average)
[1452203855] HOST ALERT: Ownerslist21;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452203871] SERVICE ALERT: Ownerslist21;Uptime;OK;SOFT;2;System Uptime - 26 day(s) 23 hour(s) 13 minute(s)
[1452203991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452204101] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452204591] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452204613] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452204991] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452205131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452205213] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452205333] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452205791] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452205911] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452205933] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452206031] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452206053] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452206161] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452206164] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452206191] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452206207] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452206228] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1452206271] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452206271] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452206311] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452206331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452206391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452206431] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452206501] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452206621] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452206653] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452206741] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452206744] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452206773] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452206807] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452206851] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 13 hour(s) 15 minute(s)
[1452206855] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1452206855] HOST FLAPPING ALERT: Ownerslist15;STARTED; Host appears to have started flapping (20.8% change > 20.0% threshold)
[1452206893] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452206931] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452206931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452206934] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452206934] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452206998] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452206998] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (21.1% change > 20.0% threshold)
[1452207031] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452207031] SERVICE FLAPPING ALERT: Ownerslist22;NSClient++ Version;STARTED; Service appears to have started flapping (22.8% change >= 20.0% threshold)
[1452207041] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.62 Gb (49%) - free 37.37 Gb (51%)
[1452207052] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2583.99 MB (42%) - free: 3530.36 MB (58%)
[1452207132] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452207135] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452207139] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452207242] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 3% (5 min average)
[1452207246] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452207256] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452207351] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452207354] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452207397] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452207418] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.01 ms
[1452207455] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452207461] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 14 hour(s) 26 minute(s)
[1452207531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452207831] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452207951] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452207997] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452208007] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452208055] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452208065] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1452208071] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452208231] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452208241] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452208301] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452208351] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452208361] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452208391] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.04 seconds)
[1452208445] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1452208471] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;3;c:\ - total: 148.47 Gb - used: 27.39 Gb (18%) - free 121.08 Gb (82%)
[1452208607] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452208665] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452208671] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452208757] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452208791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452208791] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452208820] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452208851] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452208854] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452208854] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452208884] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1452208895] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452208901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452208951] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452208961] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.64 Gb (50%) - free 37.35 Gb (50%)
[1452208971] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2580.37 MB (42%) - free: 3533.98 MB (58%)
[1452209052] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452209063] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452209071] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452209162] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452209173] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452209271] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452209274] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452209295] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452209381] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 14 hour(s) 58 minute(s)
[1452209391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452209581] SERVICE FLAPPING ALERT: Ownerslist22;Memory Usage;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1452209671] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452209791] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452209901] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452209904] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452209911] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;3;(No output returned from plugin)
[1452209968] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 3.07 ms
[1452210002] SERVICE ALERT: Ownerslist18;CPU Load;WARNING;SOFT;1;CPU Load 88% (5 min average)
[1452210011] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452210081] HOST FLAPPING ALERT: Ownerslist15;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1452210104] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452210122] SERVICE ALERT: Ownerslist18;CPU Load;WARNING;SOFT;2;CPU Load 88% (5 min average)
[1452210165] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452210242] SERVICE ALERT: Ownerslist18;CPU Load;OK;SOFT;3;CPU Load 13% (5 min average)
[1452210704] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452210765] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.71 ms
[1452211371] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452211374] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452211374] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452211481] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.66 Gb (50%) - free 37.33 Gb (50%)
[1452211485] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452211491] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2576.47 MB (42%) - free: 3537.88 MB (58%)
[1452211572] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452211575] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452211575] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452211639] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.86 ms
[1452211682] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452211693] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452211791] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452211794] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452211815] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.33 ms
[1452211901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452211901] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 15 hour(s) 40 minute(s)
[1452212207] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452212271] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452212421] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452212424] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452212488] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452212501] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452212531] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452212818] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452212861] SERVICE ALERT: Ownerslist21;W35VC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1452212882] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452213451] SERVICE ALERT: Ownerslist21;W35VC;UNKNOWN;HARD;3;Failed to open service W3SVC: 1060: The specified service does not exist as an installed service.
[1452213891] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452214001] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.56 Gb (49%) - free 37.43 Gb (51%)
[1452214001] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452214092] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452214095] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452214095] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452214111] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452214131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452214159] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.97 ms
[1452214202] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452214213] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452214231] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452214304] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452214304] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452214335] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.58 ms
[1452214335] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.8% change > 20.0% threshold)
[1452214351] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452214421] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 16 hour(s) 22 minute(s)
[1452214501] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452214504] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452214568] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452214597] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452214611] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2578.38 MB (42%) - free: 3535.97 MB (58%)
[1452214655] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.83 ms
[1452214791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452214901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452214904] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452214934] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452214937] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452215001] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452215051] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452215051] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452215171] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452215174] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452215237] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452215254] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452215257] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452215275] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452215275] HOST FLAPPING ALERT: Ownerslist15;STARTED; Host appears to have started flapping (20.1% change > 20.0% threshold)
[1452215281] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452215331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452215331] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452215381] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452215384] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452215395] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.74 ms
[1452215447] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452215485] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452215491] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 15 hour(s) 39 minute(s)
[1452215501] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452215661] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452215664] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452215685] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452215781] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452215784] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452215807] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452215807] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452215814] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452215817] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452215847] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452215875] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452215881] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452215891] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1452215921] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.58 Gb (49%) - free 37.41 Gb (51%)
[1452215931] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2578.07 MB (42%) - free: 3536.28 MB (58%)
[1452215991] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452216011] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452216023] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452216111] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452216121] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452216133] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452216231] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452216231] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452216234] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452216255] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.56 ms
[1452216281] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452216284] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452216284] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452216341] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 16 hour(s) 54 minute(s)
[1452216348] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.87 ms
[1452216391] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.42 Gb (18%) - free 121.05 Gb (82%)
[1452216411] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452216414] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452216474] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452216474] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452216484] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452216494] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452216497] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452216521] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;3;Memory usage: total:6114.35 MB - used: 2676.32 MB (44%) - free: 3438.03 MB (56%)
[1452216525] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 2.42 ms
[1452216591] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452216601] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452216621] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452216651] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452216687] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452216731] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1452216735] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452216831] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452216861] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452216864] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452216892] SERVICE ALERT: Ownerslist21;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452216928] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1452216951] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452216971] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452217002] SERVICE ALERT: Ownerslist21;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.71 Gb - used: 39.62 Gb (27%) - free 109.09 Gb (73%)
[1452217047] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452217071] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452217093] SERVICE FLAPPING ALERT: Ownerslist22;CPU Load;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1452217111] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452217191] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452217204] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452217207] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452217240] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452217270] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452217287] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452217294] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452217297] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.73 ms
[1452217321] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452217325] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1452217351] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452217421] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452217424] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452217487] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452217531] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 16 hour(s) 13 minute(s)
[1452217535] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452217791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452217847] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452217887] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452217904] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452217924] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452217928] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.76 ms
[1452217930] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 28%, RTA = 118.47 ms
[1452217951] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452218051] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452218054] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452218071] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452218117] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452218134] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452218138] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 75.53 ms
[1452218161] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;3;W3SVC: Started
[1452218191] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452218261] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452218264] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452218271] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452218324] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452218327] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452218331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452218331] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452218334] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452218334] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452218371] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;3;System Uptime - 18 day(s) 16 hour(s) 27 minute(s)
[1452218375] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452218398] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452218441] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.60 Gb (49%) - free 37.39 Gb (51%)
[1452218441] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2650.65 MB (43%) - free: 3463.70 MB (57%)
[1452218451] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2585.70 MB (42%) - free: 3528.65 MB (58%)
[1452218494] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452218504] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452218524] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452218527] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452218533] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452218537] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452218641] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 1% (5 min average)
[1452218645] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.15 ms
[1452218650] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452218751] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452218754] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452218775] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.11 ms
[1452218849] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452218861] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 17 hour(s) 36 minute(s)
[1452218913] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 602.56 ms
[1452218931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452219051] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452219054] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452219087] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452219135] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452219161] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2655.88 MB (43%) - free: 3458.47 MB (57%)
[1452219381] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452219384] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452219447] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452219448] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1452219471] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452219491] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452219511] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1452219591] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452219601] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452219631] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452219685] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452219711] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452219711] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452219991] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452219997] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452220001] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452220004] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452220007] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452220065] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452220067] HOST ALERT: Ownerslist15;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.252)
[1452220111] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.33 Gb (18%) - free 121.14 Gb (82%)
[1452220115] HOST ALERT: Ownerslist15;UP;SOFT;4;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452220251] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452220254] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452220254] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452220295] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.04 ms
[1452220301] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452220361] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.62 Gb (49%) - free 37.37 Gb (51%)
[1452220371] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2586.69 MB (42%) - free: 3527.66 MB (58%)
[1452220451] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452220453] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452220453] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452220561] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452220565] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1452220566] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452220591] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452220671] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452220674] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452220695] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452220711] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452220731] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452220781] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 18 hour(s) 8 minute(s)
[1452220791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452220894] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452220901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452220911] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452220911] SERVICE FLAPPING ALERT: Ownerslist22;Uptime;STARTED; Service appears to have started flapping (23.4% change >= 20.0% threshold)
[1452220914] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452220918] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 593.20 ms
[1452221041] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452221044] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452221107] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452221151] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;3;CPU Load 51% (5 min average)
[1452221155] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452221301] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452221311] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452221321] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452221324] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452221375] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452221411] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452221431] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.34 Gb (18%) - free 121.13 Gb (82%)
[1452221431] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452221494] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452221501] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452221504] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452221511] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452221535] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.76 ms
[1452221565] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1452221631] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452221751] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452221783] HOST ALERT: Ownerslist13;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.84)
[1452221783] SERVICE ALERT: Ownerslist13;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.84 and port 12489: No route to host
[1452221865] HOST ALERT: Ownerslist15;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452221874] SERVICE ALERT: Ownerslist13;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.84 and port 12489: No route to host
[1452221877] HOST ALERT: Ownerslist13;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.84)
[1452221894] SERVICE ALERT: Ownerslist13;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.84 and port 12489: No route to host
[1452221897] HOST ALERT: Ownerslist13;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.84)
[1452221906] SERVICE ALERT: Ownerslist13;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.84 and port 12489: No route to host
[1452221928] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452221954] SERVICE ALERT: Ownerslist13;Explorer;CRITICAL;SOFT;1;connect to address 192.168.0.84 and port 12489: No route to host
[1452221957] HOST ALERT: Ownerslist13;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.0.84)
[1452221975] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452221994] SERVICE ALERT: Ownerslist13;CPU Load;CRITICAL;SOFT;2;connect to address 192.168.0.84 and port 12489: No route to host
[1452221997] HOST ALERT: Ownerslist13;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.0.84)
[1452222014] SERVICE ALERT: Ownerslist13;W35VC;CRITICAL;SOFT;2;connect to address 192.168.0.84 and port 12489: No route to host
[1452222017] HOST ALERT: Ownerslist13;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.0.84)
[1452222026] SERVICE ALERT: Ownerslist13;C:\ Drive Space;CRITICAL;HARD;3;connect to address 192.168.0.84 and port 12489: No route to host
[1452222044] SERVICE ALERT: Ownerslist13;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.84 and port 12489: No route to host
[1452222047] HOST ALERT: Ownerslist13;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.0.84)
[1452222074] SERVICE ALERT: Ownerslist13;Explorer;CRITICAL;SOFT;2;connect to address 192.168.0.84 and port 12489: No route to host
[1452222077] HOST ALERT: Ownerslist13;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.0.84)
[1452222094] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452222101] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452222104] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452222114] SERVICE ALERT: Ownerslist13;CPU Load;CRITICAL;HARD;3;connect to address 192.168.0.84 and port 12489: No route to host
[1452222117] HOST ALERT: Ownerslist13;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.0.84)
[1452222126] SERVICE ALERT: Ownerslist13;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.84 and port 12489: No route to host
[1452222134] SERVICE ALERT: Ownerslist13;W35VC;CRITICAL;HARD;3;connect to address 192.168.0.84 and port 12489: No route to host
[1452222137] HOST ALERT: Ownerslist13;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.0.84)
[1452222137] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452222164] SERVICE ALERT: Ownerslist13;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.84 and port 12489: No route to host
[1452222165] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452222194] SERVICE ALERT: Ownerslist13;Explorer;CRITICAL;HARD;3;connect to address 192.168.0.84 and port 12489: No route to host
[1452222214] SERVICE ALERT: Ownerslist13;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.84 and port 12489: No route to host
[1452222246] SERVICE ALERT: Ownerslist13;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.84 and port 12489: No route to host
[1452222284] SERVICE ALERT: Ownerslist13;Memory Usage;CRITICAL;HARD;3;connect to address 192.168.0.84 and port 12489: No route to host
[1452222334] SERVICE ALERT: Ownerslist13;Uptime;CRITICAL;SOFT;2;connect to address 192.168.0.84 and port 12489: No route to host
[1452222366] SERVICE ALERT: Ownerslist13;NSClient++ Version;CRITICAL;HARD;3;connect to address 192.168.0.84 and port 12489: No route to host
[1452222449] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452222454] SERVICE ALERT: Ownerslist13;Uptime;CRITICAL;HARD;3;connect to address 192.168.0.84 and port 12489: No route to host
[1452222493] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452222512] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452222575] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452222694] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452222735] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452222771] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452222774] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452222774] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452222881] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.64 Gb (50%) - free 37.35 Gb (50%)
[1452222885] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452222891] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2587.73 MB (42%) - free: 3526.62 MB (58%)
[1452222951] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452222971] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452222972] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452222972] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452223047] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452223071] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452223081] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452223085] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452223086] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452223093] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452223110] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452223174] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1452223191] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452223191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452223191] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452223194] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452223213] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452223214] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452223217] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452223281] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452223287] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452223301] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 18 hour(s) 50 minute(s)
[1452223331] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452223333] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452223334] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452223337] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452223365] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452223461] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;2;CRITICAL - Socket timeout after 10 seconds
[1452223464] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452223527] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452223555] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1452223571] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;3;NSClient++ 0.4.4.15 2015-11-25
[1452223691] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452223692] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452223692] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452223801] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452223805] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1452223806] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452223904] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452223904] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452223935] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1452224021] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 19 hour(s) 2 minute(s)
[1452224091] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452224094] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452224094] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452224201] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.65 Gb (50%) - free 37.34 Gb (50%)
[1452224205] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452224211] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2589.87 MB (42%) - free: 3524.48 MB (58%)
[1452224233] HOST ALERT: Customer-Service;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.218)
[1452224274] SERVICE ALERT: Customer-Service;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.218 and port 12489: No route to host
[1452224277] HOST ALERT: Customer-Service;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.218)
[1452224340] HOST ALERT: Customer-Service;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.218)
[1452224354] SERVICE ALERT: Customer-Service;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.218 and port 12489: No route to host
[1452224357] HOST ALERT: Customer-Service;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.0.218)
[1452224394] SERVICE ALERT: Customer-Service;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.218 and port 12489: No route to host
[1452224397] HOST ALERT: Customer-Service;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.0.218)
[1452224431] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452224460] HOST ALERT: Customer-Service;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.0.218)
[1452224474] SERVICE ALERT: Customer-Service;Uptime;CRITICAL;SOFT;2;connect to address 192.168.0.218 and port 12489: No route to host
[1452224477] HOST ALERT: Customer-Service;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.0.218)
[1452224504] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452224514] SERVICE ALERT: Customer-Service;NSClient++ Version;CRITICAL;HARD;3;connect to address 192.168.0.218 and port 12489: No route to host
[1452224517] HOST ALERT: Customer-Service;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.0.218)
[1452224534] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452224534] SERVICE ALERT: Customer-Service;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.218 and port 12489: No route to host
[1452224537] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452224537] HOST ALERT: Customer-Service;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.0.218)
[1452224551] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452224561] HOST FLAPPING ALERT: Ownerslist15;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1452224594] SERVICE ALERT: Customer-Service;Uptime;CRITICAL;HARD;3;connect to address 192.168.0.218 and port 12489: No route to host
[1452224597] HOST ALERT: Customer-Service;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.0.218)
[1452224601] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452224614] SERVICE ALERT: Customer-Service;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.218 and port 12489: No route to host
[1452224651] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452224654] SERVICE ALERT: Customer-Service;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.218 and port 12489: No route to host
[1452224671] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452224704] SERVICE ALERT: Customer-Service;Explorer;CRITICAL;SOFT;1;connect to address 192.168.0.218 and port 12489: No route to host
[1452224734] SERVICE ALERT: Customer-Service;CPU Load;CRITICAL;SOFT;2;connect to address 192.168.0.218 and port 12489: No route to host
[1452224774] SERVICE ALERT: Customer-Service;C:\ Drive Space;CRITICAL;HARD;3;connect to address 192.168.0.218 and port 12489: No route to host
[1452224787] SERVICE ALERT: Customer-Service;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.218 and port 12489: No route to host
[1452224824] SERVICE ALERT: Customer-Service;Explorer;CRITICAL;SOFT;2;connect to address 192.168.0.218 and port 12489: No route to host
[1452224854] SERVICE ALERT: Customer-Service;CPU Load;CRITICAL;HARD;3;connect to address 192.168.0.218 and port 12489: No route to host
[1452224907] SERVICE ALERT: Customer-Service;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.218 and port 12489: No route to host
[1452224931] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1452224931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452224944] SERVICE ALERT: Customer-Service;Explorer;CRITICAL;HARD;3;connect to address 192.168.0.218 and port 12489: No route to host
[1452224991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452224995] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.89 ms
[1452225027] SERVICE ALERT: Customer-Service;Memory Usage;CRITICAL;HARD;3;connect to address 192.168.0.218 and port 12489: No route to host
[1452225591] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452225701] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452226007] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452226007] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452226014] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452226027] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 64%, RTA = 2.93 ms
[1452226121] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.67 Gb (50%) - free 37.32 Gb (50%)
[1452226131] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2586.36 MB (42%) - free: 3527.99 MB (58%)
[1452226191] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452226211] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452226212] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452226212] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452226295] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1452226321] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452226326] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452226431] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452226434] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452226455] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452226541] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 19 hour(s) 44 minute(s)
[1452227061] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452227064] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452227128] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.95 ms
[1452227171] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452227281] SERVICE ALERT: Ownerslist21;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452227331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452227391] SERVICE ALERT: Ownerslist21;Uptime;OK;SOFT;2;System Uptime - 27 day(s) 5 hour(s) 45 minute(s)
[1452227391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452227458] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452227498] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.80 ms
[1452227501] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452227671] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452227931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452227991] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452228104] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452228128] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.74 ms
[1452228531] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452228534] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452228534] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452228591] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452228641] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.57 Gb (49%) - free 37.42 Gb (51%)
[1452228645] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1452228651] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2592.52 MB (42%) - free: 3521.83 MB (58%)
[1452228692] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452228704] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452228724] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452228727] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452228727] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452228756] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452228841] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452228845] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452228846] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452228951] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452228954] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452228975] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.84 ms
[1452229061] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 20 hour(s) 26 minute(s)
[1452229086] HOST ALERT: Ownerslist15;DOWN;SOFT;1;PING CRITICAL - Packet loss = 100%
[1452229131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452229150] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.95 ms
[1452229150] HOST FLAPPING ALERT: Ownerslist15;STARTED; Host appears to have started flapping (23.2% change > 20.0% threshold)
[1452229241] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452229244] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452229287] HOST ALERT: Ownerslist15;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.252)
[1452229335] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452229351] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.38 Gb (18%) - free 121.10 Gb (82%)
[1452229351] HOST ALERT: Ownerslist15;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452229365] HOST ALERT: Ownerslist29;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452229471] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452229581] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452229584] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452229591] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452229648] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.87 ms
[1452229663] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452229691] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452229727] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1452229731] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452229731] SERVICE ALERT: new-pc;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452229741] SERVICE ALERT: new-pc;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452229791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452229831] SERVICE ALERT: new-pc;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452229844] SERVICE ALERT: new-pc;CPU Load;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452229854] SERVICE ALERT: new-pc;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452229888] SERVICE ALERT: new-pc;Explorer;CRITICAL;SOFT;1;connect to address 192.168.0.205 and port 12489: No route to host
[1452229944] SERVICE ALERT: new-pc;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452229961] SERVICE ALERT: Ownerslist15;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452229964] SERVICE ALERT: new-pc;CPU Load;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452229964] SERVICE ALERT: Ownerslist15;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452229964] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452229974] SERVICE ALERT: new-pc;Memory Usage;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452230008] SERVICE ALERT: new-pc;Explorer;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452230021] SERVICE ALERT: Ownerslist29;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452230044] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452230064] SERVICE ALERT: new-pc;NSClient++ Version;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452230064] SERVICE ALERT: Ownerslist29;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.144 and port 12489: No route to host
[1452230067] HOST ALERT: Ownerslist29;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.144)
[1452230071] SERVICE ALERT: Ownerslist15;C:\ Drive Space;OK;SOFT;2;c:\ - total: 148.47 Gb - used: 27.38 Gb (18%) - free 121.09 Gb (82%)
[1452230075] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452230081] SERVICE ALERT: Ownerslist15;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2660.43 MB (44%) - free: 3453.92 MB (56%)
[1452230107] HOST ALERT: Ownerslist29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452230128] SERVICE ALERT: new-pc;Explorer;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452230131] SERVICE ALERT: Ownerslist29;Memory Usage;OK;SOFT;2;Memory usage: total:9194.77 MB - used: 2860.74 MB (31%) - free: 6334.03 MB (69%)
[1452230144] SERVICE ALERT: new-pc;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.205 and port 12489: No route to host
[1452230161] SERVICE ALERT: Ownerslist15;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452230164] SERVICE ALERT: Ownerslist15;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452230164] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452230174] SERVICE ALERT: Ownerslist15;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.252 and port 12489: No route to host
[1452230181] SERVICE ALERT: Ownerslist29;Uptime;OK;SOFT;2;System Uptime - 7 day(s) 12 hour(s) 58 minute(s)
[1452230207] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452230264] SERVICE ALERT: new-pc;Uptime;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452230265] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452230271] SERVICE ALERT: Ownerslist15;CPU Load;OK;SOFT;2;CPU Load 50% (5 min average)
[1452230275] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452230281] SERVICE ALERT: Ownerslist15;W35VC;OK;SOFT;2;W3SVC: Started
[1452230291] SERVICE ALERT: Ownerslist15;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452230314] SERVICE ALERT: new-pc;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.205 and port 12489: No route to host
[1452230381] SERVICE ALERT: Ownerslist15;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452230384] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452230384] SERVICE ALERT: new-pc;Uptime;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452230419] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452230434] SERVICE ALERT: new-pc;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.205 and port 12489: No route to host
[1452230448] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452230482] HOST ALERT: Ownerslist29;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.144)
[1452230491] SERVICE ALERT: Ownerslist15;Uptime;OK;SOFT;2;System Uptime - 18 day(s) 19 hour(s) 49 minute(s)
[1452230501] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452230526] HOST ALERT: Ownerslist29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.96 ms
[1452230526] HOST FLAPPING ALERT: Ownerslist29;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1452230554] SERVICE ALERT: new-pc;C:\ Drive Space;CRITICAL;HARD;3;connect to address 192.168.0.205 and port 12489: No route to host
[1452230797] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452230807] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452230861] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1452230865] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452231051] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452231054] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452231054] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452231095] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452231101] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452231161] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.59 Gb (49%) - free 37.40 Gb (51%)
[1452231171] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 2590.55 MB (42%) - free: 3523.80 MB (58%)
[1452231251] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452231254] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452231254] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452231361] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452231365] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452231366] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452231381] HOST FLAPPING ALERT: Ownerslist29;STOPPED; Host appears to have stopped flapping (4.0% change < 5.0% threshold)
[1452231391] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452231397] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452231461] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452231471] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452231474] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452231495] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1452231511] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452231581] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 19 day(s) 21 hour(s) 8 minute(s)
[1452231591] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452231591] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452231631] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452231711] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452231831] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452231991] SERVICE ALERT: Ownerslist29;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452231994] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452231997] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452232057] HOST ALERT: Ownerslist29;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.144)
[1452232061] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452232101] SERVICE ALERT: Ownerslist29;Uptime;OK;SOFT;2;System Uptime - 7 day(s) 13 hour(s) 30 minute(s)
[1452232101] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452232104] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452232105] HOST ALERT: Ownerslist29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452232168] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452232211] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452232301] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452232304] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452232365] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452232597] HOST ALERT: Ownerslist15;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.252)
[1452232661] HOST ALERT: Ownerslist15;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452232791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server