[1452405717] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452405727] HOST ALERT: Ownerslist29;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.144)
[1452405791] HOST ALERT: Ownerslist29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452406022] SERVICE ALERT: Ownerslist29;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452406025] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452406088] HOST ALERT: Ownerslist29;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.144)
[1452406132] SERVICE ALERT: Ownerslist29;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452406136] HOST ALERT: Ownerslist29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452406191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452406304] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452406304] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452406334] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452406365] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452406398] HOST ALERT: Ownerslist29;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452406421] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 21 day(s) 21 hour(s) 42 minute(s)
[1452406471] SERVICE FLAPPING ALERT: Ownerslist22;NSClient++ Version;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1452406511] SERVICE ALERT: Ownerslist29;W35VC;CRITICAL;HARD;3;CRITICAL - Socket timeout after 10 seconds
[1452406514] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452406577] HOST ALERT: Ownerslist29;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.144)
[1452406605] HOST ALERT: Ownerslist29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452406831] SERVICE ALERT: Ownerslist29;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452406854] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452406905] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452406917] HOST ALERT: Ownerslist29;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.144)
[1452406941] SERVICE ALERT: Ownerslist29;Uptime;OK;SOFT;2;System Uptime - 9 day(s) 14 hour(s) 4 minute(s)
[1452406945] HOST ALERT: Ownerslist29;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 3.64 ms
[1452406965] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.37 ms
[1452407114] HOST ALERT: Ownerslist29;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.144)
[1452407171] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452407178] HOST ALERT: Ownerslist29;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452407291] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452407331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452407571] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452407574] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452407574] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452407581] SERVICE FLAPPING ALERT: Ownerslist22;Memory Usage;STOPPED; Service appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452407625] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452407681] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.60 Gb (49%) - free 37.39 Gb (51%)
[1452407691] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1671.53 MB (27%) - free: 4442.82 MB (73%)
[1452407701] SERVICE ALERT: Ownerslist29;W35VC;UNKNOWN;HARD;3;Failed to open service W3SVC: 1060: The specified service does not exist as an installed service.
[1452407771] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452407771] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452407774] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452407774] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452407881] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452407881] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452407885] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.52 ms
[1452407891] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452407891] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452408011] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452408131] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452408301] HOST FLAPPING ALERT: Ownerslist29;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452408408] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452408472] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1452408591] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452408831] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452408834] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452408885] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1452408941] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 21 day(s) 22 hour(s) 24 minute(s)
[1452409331] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452410091] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452410091] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452410094] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452410094] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452410145] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.87 ms
[1452410201] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.63 Gb (50%) - free 37.36 Gb (50%)
[1452410211] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1672.48 MB (27%) - free: 4441.87 MB (73%)
[1452410291] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452410291] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452410294] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452410294] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452410391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452410401] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452410401] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452410405] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.66 ms
[1452410405] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1452410411] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452410501] SERVICE FLAPPING ALERT: Ownerslist22;Explorer;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452410505] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452410569] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.87 ms
[1452411105] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452411169] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1452411351] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452411354] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452411405] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452411461] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 21 day(s) 23 hour(s) 6 minute(s)
[1452411531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452411731] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452411851] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452411871] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452412191] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452412301] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452412451] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452412471] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452412571] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452412591] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452412611] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452412621] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452412621] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452412626] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452412665] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452412691] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1452412721] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.65 Gb (50%) - free 37.34 Gb (50%)
[1452412731] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1676.10 MB (27%) - free: 4438.25 MB (73%)
[1452412811] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452412811] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452412814] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452412814] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452412896] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1452412896] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.4% change > 20.0% threshold)
[1452412901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452412921] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452412921] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452412931] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452413208] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452413265] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452413331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452413871] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452413874] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452413891] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452413925] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452413981] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 21 day(s) 23 hour(s) 48 minute(s)
[1452414011] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452414131] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452414705] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452414731] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452414731] SERVICE FLAPPING ALERT: Ownerslist22;CPU Load;STARTED; Service appears to have started flapping (22.5% change >= 20.0% threshold)
[1452414731] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 4.74 ms
[1452415131] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452415134] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452415134] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452415185] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452415191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452415241] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.67 Gb (50%) - free 37.32 Gb (50%)
[1452415251] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1673.79 MB (27%) - free: 4440.56 MB (73%)
[1452415305] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452415324] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452415324] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452415327] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452415334] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452415381] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452415441] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452415441] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452415445] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 3.48 ms
[1452415451] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452415591] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452415711] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452415753] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452415785] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452415931] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452415981] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452416051] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452416331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452416391] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452416394] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452416445] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452416501] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 0 hour(s) 30 minute(s)
[1452416651] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452416771] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452416808] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452416821] SERVICE ALERT: Ownerslist19;Memory Usage;WARNING;SOFT;1;Memory usage: total:4027.99 MB - used: 3244.74 MB (81%) - free: 783.25 MB (19%)
[1452416872] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.86 ms
[1452416891] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452416911] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452416941] SERVICE ALERT: Ownerslist19;Memory Usage;WARNING;SOFT;2;Memory usage: total:4027.99 MB - used: 3250.82 MB (81%) - free: 777.17 MB (19%)
[1452416991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452417031] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;2;(No output returned from plugin)
[1452417061] SERVICE ALERT: Ownerslist19;Memory Usage;WARNING;HARD;3;Memory usage: total:4027.99 MB - used: 3271.93 MB (81%) - free: 756.06 MB (19%)
[1452417181] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452417231] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452417408] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452417472] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452417531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452417651] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452417654] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452417654] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452417696] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452417761] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.57 Gb (49%) - free 37.42 Gb (51%)
[1452417771] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1669.98 MB (27%) - free: 4444.37 MB (73%)
[1452417781] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452417781] SERVICE FLAPPING ALERT: Ownerslist22;Memory Usage;STARTED; Service appears to have started flapping (23.9% change >= 20.0% threshold)
[1452417831] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452417851] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452417851] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452417854] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452417854] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452417918] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452417951] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452417961] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452417961] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452417971] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452418231] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452418351] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452418471] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452418551] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452418671] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452418791] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452418905] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452418905] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452418965] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452419017] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 1 hour(s) 12 minute(s)
[1452419505] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452419565] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452419931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452419991] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452420171] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452420174] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452420174] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452420220] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452420281] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.60 Gb (49%) - free 37.39 Gb (51%)
[1452420291] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1696.49 MB (28%) - free: 4417.86 MB (72%)
[1452420371] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452420371] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452420374] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452420374] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452420481] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452420481] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452420485] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.92 ms
[1452420491] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452420531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452421008] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452421072] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1452421191] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452421426] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452421429] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452421485] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.97 ms
[1452421536] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 1 hour(s) 53 minute(s)
[1452421691] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452422331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452422391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452422691] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452422691] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452422694] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452422694] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452422740] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452422801] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.62 Gb (49%) - free 37.37 Gb (51%)
[1452422811] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1684.58 MB (28%) - free: 4429.77 MB (72%)
[1452422891] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452422891] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452422891] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452422894] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452422894] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452423001] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452423001] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452423005] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.91 ms
[1452423005] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1452423011] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452423011] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452423105] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452423169] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1452423531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452423591] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452423705] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452423769] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452423946] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452423949] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452424005] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1452424056] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 2 hour(s) 35 minute(s)
[1452424211] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452424331] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452424931] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452425051] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452425171] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1452425207] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452425207] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452425214] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452425217] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452425260] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452425321] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.64 Gb (50%) - free 37.34 Gb (50%)
[1452425331] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1686.20 MB (28%) - free: 4428.15 MB (72%)
[1452425391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452425411] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452425411] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452425414] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452425414] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452425478] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1452425521] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452425521] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452425531] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452425771] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452425807] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452425863] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452425891] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452425931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452426011] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452426466] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452426469] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452426525] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452426531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452426576] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 3 hour(s) 17 minute(s)
[1452427305] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452427331] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.96 ms
[1452427731] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452427734] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452427734] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452427780] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452427791] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452427841] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.67 Gb (50%) - free 37.32 Gb (50%)
[1452427851] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1680.71 MB (27%) - free: 4433.64 MB (73%)
[1452427905] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452427911] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452427924] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452427924] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452427927] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452427934] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452428041] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452428041] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452428045] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.22 ms
[1452428051] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452428371] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452428383] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1452428986] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452428989] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452428991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452429045] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1452429096] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 3 hour(s) 59 minute(s)
[1452429111] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452429181] SERVICE FLAPPING ALERT: Ownerslist22;Memory Usage;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452429231] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452429408] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452429472] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452429531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452430008] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452430072] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1452430131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452430251] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452430254] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452430254] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452430296] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452430361] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.76 Gb (50%) - free 37.23 Gb (50%)
[1452430371] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1695.20 MB (28%) - free: 4419.15 MB (72%)
[1452430451] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452430451] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452430454] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452430454] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452430518] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.80 ms
[1452430561] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452430561] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452430571] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452430791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452431031] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452431151] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452431271] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452431505] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452431505] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452431565] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452431616] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 4 hour(s) 41 minute(s)
[1452431991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452432105] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452432165] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1452432771] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452432774] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452432774] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452432820] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1452432881] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.66 Gb (50%) - free 37.33 Gb (50%)
[1452432891] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1697.16 MB (28%) - free: 4417.19 MB (72%)
[1452432971] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452432971] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452432974] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452432974] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452433081] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452433081] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452433085] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.97 ms
[1452433091] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452433191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452433608] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452433672] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.72 ms
[1452433791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452434026] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452434029] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452434085] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452434136] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 5 hour(s) 23 minute(s)
[1452434411] SERVICE FLAPPING ALERT: Ownerslist22;CPU Load;STOPPED; Service appears to have stopped flapping (4.4% change < 5.0% threshold)
[1452434871] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452435291] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452435291] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452435294] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452435294] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452435340] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452435401] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.62 Gb (49%) - free 37.37 Gb (51%)
[1452435411] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1702.14 MB (28%) - free: 4412.21 MB (72%)
[1452435471] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452435491] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452435491] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452435494] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452435494] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452435591] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452435601] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452435601] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452435605] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1452435605] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1452435611] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452435705] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452435711] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452435769] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.85 ms
[1452436131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452436305] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452436311] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452436369] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1452436431] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452436546] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452436549] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452436551] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452436605] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.08 ms
[1452436656] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 6 hour(s) 5 minute(s)
[1452437151] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452437271] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452437391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452437391] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452437411] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452437807] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452437807] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452437814] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452437817] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452437860] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.75 ms
[1452437921] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.64 Gb (50%) - free 37.35 Gb (50%)
[1452437931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452437931] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1702.67 MB (28%) - free: 4411.68 MB (72%)
[1452438011] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452438011] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452438011] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452438014] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452438014] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452438096] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1452438121] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452438121] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452438131] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452438131] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452438408] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452438463] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.65 ms
[1452438591] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452439066] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452439069] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452439125] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452439176] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 6 hour(s) 47 minute(s)
[1452439905] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452439932] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.67 ms
[1452440331] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452440334] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452440334] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452440380] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452440391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452440441] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.66 Gb (50%) - free 37.32 Gb (50%)
[1452440451] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1705.03 MB (28%) - free: 4409.32 MB (72%)
[1452440505] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452440524] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452440524] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452440527] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452440534] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452440641] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452440641] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452440645] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 1.05 ms
[1452440651] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452440931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452440971] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452440983] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.73 ms
[1452441131] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452441251] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452441531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452441586] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452441589] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452441591] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452441645] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.86 ms
[1452441696] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 7 hour(s) 29 minute(s)
[1452441851] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452441971] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452442008] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452442072] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.99 ms
[1452442091] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452442191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452442284] HOST ALERT: Ownerslist11;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.91)
[1452442301] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452442347] HOST ALERT: Ownerslist11;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.91)
[1452442354] SERVICE ALERT: Ownerslist11;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.91 and port 12489: No route to host
[1452442381] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452442434] SERVICE ALERT: Ownerslist11;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.91 and port 12489: No route to host
[1452442437] HOST ALERT: Ownerslist11;DOWN;SOFT;3;CRITICAL - Host Unreachable (192.168.0.91)
[1452442474] SERVICE ALERT: Ownerslist11;Memory Usage;CRITICAL;SOFT;2;connect to address 192.168.0.91 and port 12489: No route to host
[1452442477] HOST ALERT: Ownerslist11;DOWN;SOFT;4;CRITICAL - Host Unreachable (192.168.0.91)
[1452442524] SERVICE ALERT: Ownerslist11;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.91 and port 12489: No route to host
[1452442527] HOST ALERT: Ownerslist11;DOWN;SOFT;5;CRITICAL - Host Unreachable (192.168.0.91)
[1452442554] SERVICE ALERT: Ownerslist11;NSClient++ Version;CRITICAL;SOFT;2;connect to address 192.168.0.91 and port 12489: No route to host
[1452442557] HOST ALERT: Ownerslist11;DOWN;SOFT;6;CRITICAL - Host Unreachable (192.168.0.91)
[1452442594] SERVICE ALERT: Ownerslist11;Memory Usage;CRITICAL;HARD;3;connect to address 192.168.0.91 and port 12489: No route to host
[1452442597] HOST ALERT: Ownerslist11;DOWN;SOFT;7;CRITICAL - Host Unreachable (192.168.0.91)
[1452442604] SERVICE ALERT: Ownerslist11;W35VC;CRITICAL;HARD;3;connect to address 192.168.0.91 and port 12489: No route to host
[1452442608] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452442644] SERVICE ALERT: Ownerslist11;Uptime;CRITICAL;SOFT;2;connect to address 192.168.0.91 and port 12489: No route to host
[1452442647] HOST ALERT: Ownerslist11;DOWN;SOFT;8;CRITICAL - Host Unreachable (192.168.0.91)
[1452442672] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1452442674] SERVICE ALERT: Ownerslist11;NSClient++ Version;CRITICAL;HARD;3;connect to address 192.168.0.91 and port 12489: No route to host
[1452442677] HOST ALERT: Ownerslist11;DOWN;SOFT;9;CRITICAL - Host Unreachable (192.168.0.91)
[1452442691] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452442704] SERVICE ALERT: Ownerslist11;C:\ Drive Space;CRITICAL;SOFT;1;connect to address 192.168.0.91 and port 12489: No route to host
[1452442707] HOST ALERT: Ownerslist11;DOWN;HARD;10;CRITICAL - Host Unreachable (192.168.0.91)
[1452442764] SERVICE ALERT: Ownerslist11;Uptime;CRITICAL;HARD;3;connect to address 192.168.0.91 and port 12489: No route to host
[1452442784] SERVICE ALERT: Ownerslist11;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.91 and port 12489: No route to host
[1452442791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452442824] SERVICE ALERT: Ownerslist11;C:\ Drive Space;CRITICAL;SOFT;2;connect to address 192.168.0.91 and port 12489: No route to host
[1452442851] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452442854] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452442854] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452442896] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1452442901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452442904] SERVICE ALERT: Ownerslist11;CPU Load;CRITICAL;SOFT;2;connect to address 192.168.0.91 and port 12489: No route to host
[1452442944] SERVICE ALERT: Ownerslist11;C:\ Drive Space;CRITICAL;HARD;3;connect to address 192.168.0.91 and port 12489: No route to host
[1452442961] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.69 Gb (50%) - free 37.30 Gb (50%)
[1452442971] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1706.60 MB (28%) - free: 4407.74 MB (72%)
[1452442981] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452443024] SERVICE ALERT: Ownerslist11;CPU Load;CRITICAL;HARD;3;connect to address 192.168.0.91 and port 12489: No route to host
[1452443051] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452443051] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452443054] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452443054] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452443161] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452443161] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452443165] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.77 ms
[1452443171] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452443473] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452443502] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.81 ms
[1452443891] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452444011] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452444104] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452444104] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452444131] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452444131] SERVICE FLAPPING ALERT: Ownerslist22;CPU Load;STARTED; Service appears to have started flapping (23.0% change >= 20.0% threshold)
[1452444165] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.60 ms
[1452444216] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 8 hour(s) 11 minute(s)
[1452444591] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452444705] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452444765] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1452445331] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452445371] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452445374] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452445374] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452445420] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452445481] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.59 Gb (49%) - free 37.40 Gb (51%)
[1452445491] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1706.10 MB (28%) - free: 4408.24 MB (72%)
[1452445571] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452445571] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452445581] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452445601] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452445665] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452445681] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452445681] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452445691] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452446208] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452446272] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.73 ms
[1452446331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452446391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452446391] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452446511] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452446531] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452446626] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452446629] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452446651] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452446685] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452446736] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 8 hour(s) 53 minute(s)
[1452446771] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452446931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452447891] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452447891] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452447894] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452447894] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452447940] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452448001] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.61 Gb (49%) - free 37.38 Gb (51%)
[1452448011] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1708.16 MB (28%) - free: 4406.19 MB (72%)
[1452448091] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452448091] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452448094] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452448094] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452448131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452448201] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452448201] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452448205] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.04 ms
[1452448205] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1452448211] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452448305] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452448311] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452448369] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452448431] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452448905] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452448969] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1452449144] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452449148] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452449205] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.51 ms
[1452449254] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 9 hour(s) 35 minute(s)
[1452449631] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452449751] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452449871] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452449991] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452450411] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452450421] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452450421] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452450426] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452450458] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452450521] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.63 Gb (50%) - free 37.36 Gb (50%)
[1452450531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452450531] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1709.39 MB (28%) - free: 4404.96 MB (72%)
[1452450611] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452450611] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452450614] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452450614] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452450696] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452450696] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.4% change > 20.0% threshold)
[1452450721] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452450721] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452450731] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452451008] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452451061] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 2.86 ms
[1452451664] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452451667] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452451671] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452451725] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1452451731] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452451774] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 10 hour(s) 17 minute(s)
[1452452271] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452452391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452452391] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452452505] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452452511] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452452531] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.83 ms
[1452452931] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452452941] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452453041] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.65 Gb (50%) - free 37.34 Gb (50%)
[1452453051] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1711.56 MB (28%) - free: 4402.79 MB (72%)
[1452453105] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452453124] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452453124] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452453127] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452453134] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452453241] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452453241] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452453245] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.21 ms
[1452453251] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452453531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452453571] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452453581] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.68 ms
[1452453971] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452454184] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452454187] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452454191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452454245] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452454294] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 10 hour(s) 59 minute(s)
[1452454301] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452454571] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452454608] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452454672] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1452454691] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452454811] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452454901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452455208] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452455272] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.78 ms
[1452455331] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452455451] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452455454] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452455454] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452455496] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452455561] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.67 Gb (50%) - free 37.32 Gb (50%)
[1452455571] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1714.46 MB (28%) - free: 4399.89 MB (72%)
[1452455651] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452455651] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452455654] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452455654] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452455761] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452455761] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452455765] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.00 ms
[1452455771] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452456073] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452456100] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.65 ms
[1452456111] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452456231] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452456351] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452456531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452456703] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452456703] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452456765] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.02 ms
[1452456814] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 11 hour(s) 41 minute(s)
[1452456951] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452457071] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452457191] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452457305] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452457365] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.58 ms
[1452457901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452457971] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452457974] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452457974] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452458018] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452458081] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.57 Gb (49%) - free 37.42 Gb (51%)
[1452458091] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1716.64 MB (28%) - free: 4397.71 MB (72%)
[1452458171] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452458171] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452458174] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452458174] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452458281] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452458281] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452458285] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.89 ms
[1452458291] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452458501] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452458501] SERVICE FLAPPING ALERT: Ownerslist22;Explorer;STARTED; Service appears to have started flapping (23.2% change >= 20.0% threshold)
[1452458808] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452458872] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.32 ms
[1452459224] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452459227] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452459285] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.25 ms
[1452459334] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 12 hour(s) 23 minute(s)
[1452460191] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452460191] SERVICE FLAPPING ALERT: Ownerslist22;Uptime;STARTED; Service appears to have started flapping (20.5% change >= 20.0% threshold)
[1452460491] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452460491] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452460494] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452460494] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452460538] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452460601] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.60 Gb (49%) - free 37.39 Gb (51%)
[1452460611] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1717.44 MB (28%) - free: 4396.90 MB (72%)
[1452460691] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452460691] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452460694] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452460694] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452460791] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452460791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452460801] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452460801] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452460805] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452460805] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1452460811] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452460905] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452460911] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452460969] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.80 ms
[1452461391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452461505] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452461569] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452461744] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452461747] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452461805] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.53 ms
[1452461854] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 13 hour(s) 5 minute(s)
[1452461991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452462011] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452462711] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452462831] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452462951] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452463011] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452463021] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452463021] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452463026] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452463058] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1452463121] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.62 Gb (49%) - free 37.37 Gb (51%)
[1452463131] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1720.36 MB (28%) - free: 4393.99 MB (72%)
[1452463211] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452463211] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452463214] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452463214] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452463261] SERVICE ALERT: Ownerslist19;Memory Usage;OK;HARD;3;Memory usage: total:4331.99 MB - used: 3439.45 MB (79%) - free: 892.54 MB (21%)
[1452463296] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.68 ms
[1452463296] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.4% change > 20.0% threshold)
[1452463321] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452463321] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452463331] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452463608] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452463661] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.74 ms
[1452463731] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452463811] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452463931] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452464264] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452464267] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452464325] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452464374] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 13 hour(s) 47 minute(s)
[1452464391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452465105] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452465131] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.65 ms
[1452465531] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452465531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452465541] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452465641] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.64 Gb (50%) - free 37.35 Gb (50%)
[1452465651] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1715.55 MB (28%) - free: 4398.80 MB (72%)
[1452465705] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452465724] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452465724] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452465727] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452465734] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452465841] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452465841] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452465845] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 0.56 ms
[1452465851] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452466157] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452466181] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 2.01 ms
[1452466784] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452466787] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452466791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452466845] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1452466894] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 14 hour(s) 29 minute(s)
[1452467151] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452467208] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452467272] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.77 ms
[1452467671] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452467751] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452467808] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452467871] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452467872] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.88 ms
[1452467991] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452468051] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452468054] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452468054] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452468096] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.62 ms
[1452468131] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452468161] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.67 Gb (50%) - free 37.32 Gb (50%)
[1452468171] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1769.65 MB (29%) - free: 4344.70 MB (71%)
[1452468251] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452468251] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452468251] SERVICE FLAPPING ALERT: Ownerslist22;CPU Load;STOPPED; Service appears to have stopped flapping (4.9% change < 5.0% threshold)
[1452468251] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452468254] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452468254] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452468361] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452468361] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452468365] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.65 ms
[1452468371] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452468591] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452468673] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452468699] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.75 ms
[1452469191] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452469304] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452469304] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452469365] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.63 ms
[1452469414] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 15 hour(s) 11 minute(s)
[1452469471] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452469591] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452469711] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452469901] SERVICE FLAPPING ALERT: Ownerslist22;Explorer;STOPPED; Service appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452469905] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452469965] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452470391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452470571] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452470574] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452470574] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452470618] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1452470681] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.69 Gb (50%) - free 37.30 Gb (50%)
[1452470691] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1769.68 MB (29%) - free: 4344.67 MB (71%)
[1452470771] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452470771] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452470781] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452470801] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1452470865] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.97 ms
[1452470881] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452470881] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452470891] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452470991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452471408] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452471472] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.83 ms
[1452471824] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452471827] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452471885] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.14 ms
[1452471934] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 15 hour(s) 53 minute(s)
[1452473091] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452473091] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452473094] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452473094] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452473138] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452473201] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.59 Gb (49%) - free 37.40 Gb (51%)
[1452473211] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1771.61 MB (29%) - free: 4342.74 MB (71%)
[1452473291] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452473291] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452473294] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452473294] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452473391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452473401] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452473401] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452473405] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452473405] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.0% change > 20.0% threshold)
[1452473411] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452473505] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452473569] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.79 ms
[1452473931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452473991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452474105] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452474169] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.69 ms
[1452474251] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452474344] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452474347] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452474371] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452474405] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 4.05 ms
[1452474454] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 16 hour(s) 35 minute(s)
[1452474491] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452475191] SERVICE FLAPPING ALERT: Ownerslist22;Uptime;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1452475611] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452475621] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452475621] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.9% change < 5.0% threshold)
[1452475626] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452475658] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.63 ms
[1452475721] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.61 Gb (49%) - free 37.38 Gb (51%)
[1452475731] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1772.38 MB (29%) - free: 4341.97 MB (71%)
[1452475791] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452475811] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452475811] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452475814] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452475814] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452475896] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.54 ms
[1452475896] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.4% change > 20.0% threshold)
[1452475901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452475921] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452475921] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452475931] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452476226] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452476260] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.76 ms
[1452476291] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452476311] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452476391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452476864] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452476867] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452476911] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452476925] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.67 ms
[1452476931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452476974] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 17 hour(s) 17 minute(s)
[1452476991] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452477031] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452477151] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452477151] SERVICE FLAPPING ALERT: Ownerslist22;NSClient++ Version;STARTED; Service appears to have started flapping (21.3% change >= 20.0% threshold)
[1452477701] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452477705] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452477731] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 54%, RTA = 0.86 ms
[1452478091] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452478131] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452478131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452478141] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452478211] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452478241] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.64 Gb (50%) - free 37.35 Gb (50%)
[1452478251] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1779.38 MB (29%) - free: 4334.97 MB (71%)
[1452478301] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452478305] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452478324] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452478324] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452478327] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452478331] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;HARD;3;could not fetch information from server
[1452478334] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452478351] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;HARD;3;(No output returned from plugin)
[1452478441] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452478441] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452478445] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.06 ms
[1452478451] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452478771] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452478780] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 28%, RTA = 0.70 ms
[1452478951] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452479071] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452479191] SERVICE ALERT: Ownerslist22;NSClient++ Version;OK;SOFT;3;(No output returned from plugin)
[1452479384] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452479387] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452479391] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452479445] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 2.64 ms
[1452479494] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 17 hour(s) 59 minute(s)
[1452479501] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452479501] SERVICE FLAPPING ALERT: Ownerslist22;Explorer;STARTED; Service appears to have started flapping (23.6% change >= 20.0% threshold)
[1452479808] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452479872] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.28 ms
[1452479931] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452480391] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;1;could not fetch information from server
[1452480426] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.03 seconds)
[1452480490] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.88 ms
[1452480511] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;SOFT;2;could not fetch information from server
[1452480591] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452480631] SERVICE ALERT: Ownerslist22;NSClient++ Version;WARNING;HARD;3;could not fetch information from server
[1452480651] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452480654] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452480654] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452480696] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1452480701] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452480761] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.66 Gb (50%) - free 37.33 Gb (50%)
[1452480771] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1777.92 MB (29%) - free: 4336.43 MB (71%)
[1452480781] SERVICE ALERT: Ownerslist22;Memory Usage;UNKNOWN;HARD;3;could not fetch information from server
[1452480851] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452480851] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452480854] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452480854] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452480961] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452480961] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452480965] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.76 ms
[1452480971] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452481191] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452481273] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452481300] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 44%, RTA = 0.86 ms
[1452481381] SERVICE ALERT: Ownerslist22;Memory Usage;WARNING;HARD;3;could not fetch information from server
[1452481901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452481971] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452481974] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452481974] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452482081] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.67 Gb (50%) - free 37.32 Gb (50%)
[1452482085] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.61 ms
[1452482091] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1786.12 MB (29%) - free: 4328.23 MB (71%)
[1452482171] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452482171] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452482181] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452482201] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452482265] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.05 ms
[1452482281] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452482281] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452482291] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452482391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452482501] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452482531] SERVICE ALERT: Ownerslist22;CPU Load;OK;HARD;3;CPU Load 0% (5 min average)
[1452482531] SERVICE FLAPPING ALERT: Ownerslist22;CPU Load;STARTED; Service appears to have started flapping (21.1% change >= 20.0% threshold)
[1452482826] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452482885] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.70 ms
[1452483408] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452483473] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.81 ms
[1452483591] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452483701] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452483731] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452483851] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452483891] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452483894] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452483894] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452484001] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.69 Gb (50%) - free 37.30 Gb (50%)
[1452484005] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.48 ms
[1452484011] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1786.45 MB (29%) - free: 4327.90 MB (71%)
[1452484091] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452484091] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452484094] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452484094] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452484191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452484201] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452484201] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452484205] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.64 ms
[1452484211] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452484301] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452484304] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452484304] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452484368] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.82 ms
[1452484414] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 19 hour(s) 21 minute(s)
[1452484791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452484901] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452485391] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452485391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452485511] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452485614] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452485631] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452485811] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452485821] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452485827] HOST ALERT: Ownerslist17;DOWN;SOFT;1;(Host check timed out after 30.02 seconds)
[1452485891] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.81 ms
[1452485921] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.58 Gb (49%) - free 37.41 Gb (51%)
[1452485931] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1791.70 MB (29%) - free: 4322.65 MB (71%)
[1452485991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452486011] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452486011] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452486014] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452486014] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452486096] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.57 ms
[1452486096] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.5% change > 20.0% threshold)
[1452486121] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452486121] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452486131] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452486224] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452486227] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452486231] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452486291] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 1.43 ms
[1452486334] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 19 hour(s) 53 minute(s)
[1452486351] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452486471] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452486531] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452487071] SERVICE ALERT: Ownerslist22;Uptime;OK;HARD;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452487131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;WARNING;HARD;3;could not fetch information from server
[1452487191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452487534] HOST FLAPPING ALERT: Ownerslist17;STOPPED; Host appears to have stopped flapping (3.8% change < 5.0% threshold)
[1452487791] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452487901] SERVICE ALERT: Ownerslist22;Explorer;WARNING;HARD;3;could not fetch information from server
[1452487905] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452487930] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 50%, RTA = 0.73 ms
[1452488051] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452488171] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452488271] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452488291] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1452488331] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452488334] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452488334] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452488391] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452488391] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452488441] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.61 Gb (49%) - free 37.38 Gb (51%)
[1452488445] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.60 ms
[1452488445] HOST FLAPPING ALERT: Ownerslist17;STARTED; Host appears to have started flapping (22.2% change > 20.0% threshold)
[1452488451] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1790.00 MB (29%) - free: 4324.35 MB (71%)
[1452488501] SERVICE ALERT: Ownerslist22;Explorer;UNKNOWN;HARD;3;could not fetch information from server
[1452488505] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452488524] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452488524] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452488527] HOST ALERT: Ownerslist17;DOWN;SOFT;2;CRITICAL - Host Unreachable (192.168.0.180)
[1452488534] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452488641] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452488641] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452488645] HOST ALERT: Ownerslist17;UP;SOFT;3;PING OK - Packet loss = 0%, RTA = 2.09 ms
[1452488651] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452488744] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452488747] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452488811] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.74 ms
[1452488854] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 20 hour(s) 35 minute(s)
[1452488891] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452488991] SERVICE ALERT: Ownerslist22;W35VC;WARNING;HARD;3;could not fetch information from server
[1452489011] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452489591] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452489711] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;2;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452490008] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452490058] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.59 ms
[1452490131] SERVICE ALERT: Ownerslist22;C:\ Drive Space;UNKNOWN;HARD;3;Free disk space : Invalid drive
[1452490191] SERVICE ALERT: Ownerslist22;W35VC;UNKNOWN;HARD;3;could not fetch information from server
[1452490211] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452490311] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452490331] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452490431] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452490451] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)
[1452490551] SERVICE ALERT: Ownerslist22;Uptime;OK;SOFT;3;System Uptime - 0 day(s) 0 hour(s) 0 minute(s)
[1452490608] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452490661] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 37%, RTA = 0.85 ms
[1452490831] SERVICE FLAPPING ALERT: Ownerslist22;NSClient++ Version;STOPPED; Service appears to have stopped flapping (4.0% change < 5.0% threshold)
[1452490851] SERVICE ALERT: Ownerslist17;C:\ Drive Space;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452490854] SERVICE ALERT: Ownerslist17;Memory Usage;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452490854] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452490896] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.81 ms
[1452490961] SERVICE ALERT: Ownerslist17;C:\ Drive Space;OK;SOFT;2;c:\ - total: 73.99 Gb - used: 36.63 Gb (50%) - free 37.36 Gb (50%)
[1452490971] SERVICE ALERT: Ownerslist17;Memory Usage;OK;SOFT;2;Memory usage: total:6114.35 MB - used: 1789.88 MB (29%) - free: 4324.47 MB (71%)
[1452491051] SERVICE ALERT: Ownerslist17;NSClient++ Version;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452491051] SERVICE ALERT: Ownerslist17;CPU Load;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452491051] SERVICE FLAPPING ALERT: Ownerslist22;CPU Load;STOPPED; Service appears to have stopped flapping (4.5% change < 5.0% threshold)
[1452491054] SERVICE ALERT: Ownerslist17;W35VC;CRITICAL;SOFT;1;connect to address 192.168.0.180 and port 12489: No route to host
[1452491054] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452491151] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;1;could not fetch information from server
[1452491161] SERVICE ALERT: Ownerslist17;NSClient++ Version;OK;SOFT;2;NSClient++ 0.4.4.15 2015-11-25
[1452491161] SERVICE ALERT: Ownerslist17;CPU Load;OK;SOFT;2;CPU Load 0% (5 min average)
[1452491165] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.55 ms
[1452491171] SERVICE ALERT: Ownerslist17;W35VC;OK;SOFT;2;W3SVC: Started
[1452491264] SERVICE ALERT: Ownerslist17;Uptime;CRITICAL;SOFT;1;CRITICAL - Socket timeout after 10 seconds
[1452491267] HOST ALERT: Ownerslist17;DOWN;SOFT;1;CRITICAL - Host Unreachable (192.168.0.180)
[1452491271] SERVICE ALERT: Ownerslist22;Uptime;WARNING;SOFT;2;could not fetch information from server
[1452491331] HOST ALERT: Ownerslist17;UP;SOFT;2;PING OK - Packet loss = 0%, RTA = 0.71 ms
[1452491374] SERVICE ALERT: Ownerslist17;Uptime;OK;SOFT;2;System Uptime - 22 day(s) 21 hour(s) 17 minute(s)
[1452491391] SERVICE ALERT: Ownerslist22;Uptime;WARNING;HARD;3;could not fetch information from server
[1452491651] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;1;could not fetch information from server
[1452491771] SERVICE ALERT: Ownerslist22;CPU Load;WARNING;SOFT;2;could not fetch information from server
[1452491891] SERVICE ALERT: Ownerslist22;CPU Load;OK;SOFT;3;CPU Load 0% (5 min average)